# Setup

**CS5483 Data Warehousing and Data Mining**
___

## JupyterHub

### How to access the JupyterHub server?

- Enter the url of the Jupyterhub server [divedeep.cs.cityu.edu.hk/cs5483_22b](https://divedeep.cs.cityu.edu.hk/cs5483_22b) in a web browser.
- Enter your [EID](https://www.cityu.edu.hk/esu/eid.htm) and Password in the fields `Username` and `Password` respectively.
- Click the `Sign in` button.  

![log in](images/login.dio.svg)

- Choose the desired server option and click `Start`.  

![server options](images/serveroptions.dio.svg) 

The JupyterLab interface should appear.

```{note}

- `GPU Notebook` may fail to spawn if GPU resources run out. You may use `CPU Notebook` instead.
```

```{tip}

**How to troubleshoot in case or error?**

- Refresh the page, re-login, or clear your browser cache to reset your browser.
- Restart your server using the `Hub Control Panel` under the `File` menu (shown above) or visit the hub home  
  https://divedeep.cs.cityu.edu.hk/cs5483_22b/hub/home  
  and click `Stop My Server` and then `Start My Server`.  
  ![Restart my server](images/restart.dio.svg)
- Contact your instructor using the contact information on the [course homepage](https://www.cs.cityu.edu.hk/~ccha23/cs5483).

```

## Jupyter Notebook

### How to access course notebooks?

The recommended way is to follow the links on the [course homepage](https://www.cs.cityu.edu.hk/~ccha23/cs5483).

```{important}

It is important to open the notebooks using the links on the course homepage as it 
- [git-pull](https://git-scm.com/docs/git-pull) updates/corrections to the notebooks and 
- merge them with the notebooks stored under `cs5483` folder in your home directory but
- without overwriting your changes.

If you open the notebook directly on JupyterHub instead of following the git-pull link, updates will not be pulled, and your notebook may fail to have the updates needed for submission.

```

### How to complete tutorial exercises?

Learn how to edit a notebook:


- Watch the official [video tutorial](https://youtu.be/A5YyoCKxEOU) to JupyterLab.

**Exercise** 

As a starter, write the following python program ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) to say Hello to the world. Type the program below and use tab or 4 spaces to indent. Run it with `Shift+Enter`. 

![Hello world program](images/hello.dio.svg)

<!-- grade_id: say_hello, points: 0, grade: false, solution: true -->
```python
# your python code here
# end of python code

```

```{attention}

We often ask you to write your answer/code in a special cells with *unique but hidden id number*. 
- Make sure you provide your answers only to the answer cells that says `YOUR CODE HERE` or "YOUR ANSWER HERE". 
- For coding exercises, you should also remove `raise NotImplementedError()`.
- Do not clone or duplicate the answer cells.

```

In order to check your work thoroughly, there will be both *visible* and *hidden* test cases. The following is a visible test you can run to check your answer:

<!-- grade_id: test-say_hello, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your "Hello, World!" program.
import io
import sys

old_stdout, sys.stdout = sys.stdout, io.StringIO()
try:
    say_hello()
    assert sys.stdout.getvalue() == "Hello, World!\n"
finally:
    sys.stdout = old_stdout
```

The test returns an assertion error only if your program does not print the correct message.

```{tip}

1. You can repeatedly modify your solution and run the test cell until your solution passes the test. There is no mark penalty in failing the test before submission. 
1. The test cell often uses more tools that what you might have learned. You are not required to know how to write the test cell but you should try to understand the error messages from a failed test.
1. To assess your solution thoroughly, we will run new tests hidden from you after you have submitted your notebook. There is no partial credit for a partially correct solution that works for the visible test but fails for the hidden test. Therefore, *you should ensure your solution works in general rather than just the visible tests*.
1. If you open the same notebook multiple times in different browser windows, be careful in making changes in different windows as you may overwrite your own changes.
1. If your notebook fails to run any code, the Kernel might have died. You can restart the kernel with `Kernel`$\to$ `Restart`. If restarting fails, check your code cells to see if there is any code that breaks the kernel.

```

### How to submit a notebook

To submit your notebook:

1. Go to `Nbgrader->Assignment List` and . 
1. (Optional) Expand the Lab folder and click the `validate` button next to the notebook(s) to check if all the visible tests pass.
1. Click `Submit` to submit your notebook.

```{attention}

Before you submit, make sure everything runs as expected:

1. **Git-pull the notebooks**: Follow any one of the link on the course homepage to a notebook, which will git-pull any updates/corrections to (all) your notebooks.
1. **Restart the kernel**: `Kernel`$\to$ `Restart Kernel...` 
1. **run all cells**: `Run`$\to$ `Run All Cells`

You may submit as many times as you wish before the due date as we collect your latest submission for grading. *No late submission* will be collected without valid justifications.

```

```{important}

Auto-grading often **FAILs** (after submission) if:

1. The notebook files are renamed, e.g., `Setup.ipynb` changed/copied to `Setup (Copy).ipynb`, or converted to a python script `Setup.py`.
2. An html file exists with the same name as a notebook, e.g., `Setup.html` is the default grading feedback file to be generated when grading `Setup.ipynb`.
3. The file size is too large, e.g, over `100MB`.
4. The notebook is obtained from or shared with another student.

```

```{tip}

In case of error: 

1. Download/backup your existing notebooks,
2. remove them from the `Lab*` folder,
3. click the git-pull links from the course homepage to re-pull the notebooks, and
4. copy your solutions to the new notebooks.

```